<figure>
  <IMG SRC="https://upload.wikimedia.org/wikipedia/commons/thumb/d/d5/Fachhochschule_Südwestfalen_20xx_logo.svg/320px-Fachhochschule_Südwestfalen_20xx_logo.svg.png" WIDTH=250 ALIGN="right">
</figure>

# Programmierung für KI
### Winterersemester 2024/25
Prof. Dr. Heiner Giefers

**Übung 1: Was ist der Zweck des Schlüsselworts def in Python?**

a) Es ist Slang und bedeutet „der folgende Code ist wirklich cool“ <br>
b) Es zeigt den Beginn einer Funktion an <br>
c) Es zeigt an, dass der folgende eingerückte Codeabschnitt für eine spätere Wiederverwendung abgelegt werden soll <br>
d) b und c sind beide wahr <br>
e) Keiner der oben genannten Punkte <br>

Tragen Sie ihre Antwort in den return-Wert ein.

```python
!pip install nose
```

Output:
```
Requirement already satisfied: nose in /Users/andreaseri/miniconda3/envs/grundlagen-programmieren-test/lib/python3.11/site-packages (1.3.7)
```

<!-- grade_id: cell-6d1235dc06fffa57, points: 0, grade: false, solution: true -->
```python
def answer1():
    antwort = None
    antwort = 'd'
    return antwort
```

<!-- grade_id: cell-61aeff3d56ad49f3, points: 1, grade: true, solution: false -->
```python
assert type(answer1())==str, "Bitte setzen Sie antwort auf eine Zeichenkette 'a', 'b', 'c', 'd' oder 'e'"
```

**Übung 2: Was wird das folgende Python Programm ausgeben?**

```python
def fred():
    print("Zap")
def jane():
    print("ABC")

jane()
fred()
jane()
```


a) Zap ABC jane fred jane <br>
b) Zap ABC Zap <br>
c) ABC Zap jane <br>
d) ABC Zap ABC <br>
e) Zap Zap Zap <br>

Tragen Sie ihre Antwort in den return-Wert ein.

<!-- grade_id: cell-1ab8a61d271937a1, points: 0, grade: false, solution: true -->
```python
def answer1():
    antwort = None
    antwort = 'd'
    return antwort
```

<!-- grade_id: cell-426cda653caad121, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(type(answer1()), str, "Bitte setzen Sie antwort auf eine Zeichenkette 'a', 'b', 'c', 'd' oder 'e'")
```

**Übung 3: Schreiben Sie ihr Programm zur Lohnberechnung aus Kapitel 3 so um, dass es als Funktion ausgeführt werden kann. Die Funktion soll "lohnberechnung" heißen, die Parameter `stundenlohn` und `arbeitsstunden` besitzen sowie den berechneten Gehaltswert zurück geben. Bei Übergabe einer nicht numerischen Eingabe, soll die Fehlermeldung `Fehler, bitte gib eine Zahl ein` zurückgegeben werden.**

<!-- grade_id: cell-160086955c8b2f59, points: 0, grade: false, solution: true -->
```python
def get_float(value):
    try:
        return True, float(value)
    except:
        return False, None

# 1,5 fache ab 40h
def lohnberechnung(stundenlohn, arbeitsstunden):
    for value in [stundenlohn, arbeitsstunden]:
        ok, _ = get_float(value)
        if not ok:
            return "Fehler, bitte gib eine Zahl ein."
        
    mehrarbeit_ab = 40
    gehalt = stundenlohn * (min(arbeitsstunden, mehrarbeit_ab) + 1.5 * max(arbeitsstunden - mehrarbeit_ab, 0))
    return gehalt
```

<!-- grade_id: cell-79654f2541fc33a8, points: 1, grade: true, solution: false -->
```python
"""Überprüft, dass die Funktion lohnberechnung korrekt funktioniert."""
from nose.tools import assert_almost_equal
assert_almost_equal(lohnberechnung(10, 45), 475.0, "Falsches Ergebnis")
assert_almost_equal(lohnberechnung(10, 15), 150.0, "Falsches Ergebnis")
assert_equal(lohnberechnung("zehn", 10), "Fehler, bitte gib eine Zahl ein.")
assert_equal(lohnberechnung(10, "zehn"), "Fehler, bitte gib eine Zahl ein.")
```

**Übung 4:  Schreiben Sie das Benotungsprogramm aus Kapitel 3 neu, indem Sie eine Funktion namens `notenberechnung` implementieren, die eine Punktzahl als Parameter annimmt und eine Note zurückgibt. Falls die Eingabe keine Zahl ist, geben Sie `Bitte eine Zahl eingeben` als Ergebnis zurück und beenden damit die Funktion. Geben Sie die Fehlermeldung `Falsche Punktanzahl, bitte geben Sie die Punkte von 0.0 bis 1.0 ein` als Ergebnis der Funktion zurück, falls die eingegebene Zahl größer 1 und kleiner 0 ist.**

Punkte   Note 
 >= 0.9     1  
 >= 0.8     2  
 >= 0.7     3  
 >= 0.6     4  
  < 0.6     5  

<!-- grade_id: cell-968a8a9e7e209e84, points: 0, grade: false, solution: true -->
```python
def get_float(value):
    try:
        return True, float(value)
    except:
        return False, None
        
def notenberechnung(punktzahl):
    ok, punktzahl = get_float(punktzahl)
    if not ok:
        return "Bitte eine Zahl eingeben"

    if punktzahl < 0 or punktzahl > 1:
        return "Falsche Punktanzahl, bitte geben Sie die Punkte von 0.0 bis 1.0 ein"

    if punktzahl >= 0.9:
        return 1
    elif punktzahl >= 0.8:
        return 2
    elif punktzahl >= 0.7:
        return 3
    elif punktzahl >= 0.6:
        return 4
    else:
        return 5
```

<!-- grade_id: cell-a4b1592b04a632f4, points: 1, grade: true, solution: false -->
```python
"""Überprüft, dass die Funktion notenberechnung korrekt funktioniert."""
from nose.tools import assert_equal
assert_equal(notenberechnung(0.5), 5, "Falsches Ergebnis")
assert_equal(notenberechnung(0.6), 4, "Falsches Ergebnis")
assert_equal(notenberechnung(0.7), 3, "Falsches Ergebnis")
assert_equal(notenberechnung(0.8), 2, "Falsches Ergebnis")
assert_equal(notenberechnung(0.9), 1, "Falsches Ergebnis")
assert_equal(notenberechnung(1.1), "Falsche Punktanzahl, bitte geben Sie die Punkte von 0.0 bis 1.0 ein")
assert_equal(notenberechnung(-1.0), "Falsche Punktanzahl, bitte geben Sie die Punkte von 0.0 bis 1.0 ein")
assert_equal(notenberechnung("eins"), "Bitte eine Zahl eingeben")
```